<!-- grade_id: palindrome, points: 0, grade: false, solution: true -->
```python
def is_palindrome(s):
    # YOUR CODE HERE
    for i,c in enumerate(s):
        if c != s[-i-1]:
            return False
    return True
        
```

<!-- grade_id: palindrome_tests1, points: 2, grade: true, solution: false -->
```python
s = "absdf"
assert is_palindrome(s) in (True, False)
```

<!-- grade_id: palindrome_tests2, points: 2, grade: true, solution: false -->
```python

```